// Copyright 2026 the Kessel authors.  All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-literals-numeric-literals
info: Numeric separator literal syntax parses in script goal
description: Minimal numeric separator smoke test
---*/

var decimal = 1_000_000;
var binary = 0b1010_1010;
