// Copyright 2026 the Kessel authors.  All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.

/*---
esid: sec-literals-numeric-literals
info: BigInt literal syntax parses in script goal
description: Minimal BigInt literal smoke test
---*/

var decimal = 123n;
var binary = 0b1010n;
